# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

def proportion_of_education():

    # Maternal education 
    # 1: <12 years, 
    # 2:12 years, 
    # 3: >12 years, not a college graduate
    # 4: College graduate
    columns = ['EDUC1']
    education_df = pd.read_csv("NISPUF17.csv", usecols=columns)
    df_len = len(education_df)
    data = education_df.value_counts()
    
    return {"less than high school": data.get(1.0)/df_len, "high school": data.get(2.0)/df_len, "more than high school but not college": data.get(3.0)/df_len, "college": data.get(4.0)/df_len}

    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # Was child ever fed breast milk?
    #    Yes
    #    No
    #    Don’t know
    #    Missing
    
    columns               = ['CBF_01', 'P_NUMFLU']
    bf_and_vaccine_df     = pd.read_csv("NISPUF17.csv", usecols=columns)
    df_len                = len(bf_and_vaccine_df)

    # remove NA data from breast_and_vaccine_df
    bf_and_vaccine_df.dropna(axis=0, inplace=True)

    # get entries with CBF_01 == 1 AND CBF_01 == 2
    # breast_fed is a data frame containing the 
    breast_fed     = bf_and_vaccine_df[bf_and_vaccine_df['CBF_01'] == 1]
    not_breast_fed = bf_and_vaccine_df[bf_and_vaccine_df['CBF_01'] == 2]

    # get groups with data from P_NUMFLU from each 
    bf_data  = breast_fed['CBF_01'].value_counts()
    nbf_data = not_breast_fed['CBF_01'].value_counts()

    n_bf_injected          = 0
    n_bf_total_injections  = 0
    for i in range(len(bf_data)) :
        n_bf_injected         = n_bf_injected + bf_data[1, i]
        n_bf_total_injections = n_bf_total_injections + (i * bf_data[1, i])

        print(i, "  ", bf_data[1, i], "  ", i*bf_data[1, i])

#    print(type(bf_data))
    print("BF Data        \n", bf_data)
    print("NBF Data        \n", nbf_data)
#    print("breast_fed     \n", len(breast_fed))
#    print("BF Data Index  \n", bf_data.index)
#    print("BF Data Length \n", len(bf_data))
#    print("BF Injected    \n", n_bf_injected)
#    print("BF Injections  \n", n_bf_total_injections)


    n_nbf_injected         = 0
    n_nbf_total_injections = 0
    for i in range(len(nbf_data)) :
        n_nbf_injected         = n_nbf_injected + nbf_data[2, i]
        n_nbf_total_injections = n_nbf_total_injections + (i * nbf_data[2, i])

#        print(i, "  ", nbf_data[2, i], "  ", i*nbf_data[2, i])

#    print(type(nbf_data))
#    print("NBF Data        \n", nbf_data)
#    print("not_breast_fed  \n", not_breast_fed)
#    print("NBF Data Index  \n", nbf_data.index)
#    print("NBF Data Length \n", len(nbf_data))
#    print("NBF Injected    \n", n_nbf_injected)
#    print("NBF Injections  \n", n_nbf_total_injections)

    return (n_bf_total_injections/n_bf_injected, n_nbf_total_injections/n_nbf_injected)
    raise NotImplementedError()
```

```python
print(average_influenza_doses())
#(1.8799187420058687, 1.5963945918878317)
# BF Data        
# CBF_01  P_NUMFLU
# 1       3.0         4334
#         0.0         3466
#         2.0         2917
#         1.0         1395
#         4.0         1143
#         5.0           33
#         6.0            3
# dtype: int64
# NBF Data        
#  CBF_01  P_NUMFLU
# 2       0.0         656
#         3.0         503
#         2.0         422
#         1.0         277
#         4.0         137
#         5.0           2
```

Output:
```
0    3466    0
1    1395    1395
2    2917    5834
3    4334    13002
4    1143    4572
5    33    165
6    3    18
BF Data        
 CBF_01  P_NUMFLU
1       3.0         4334
        0.0         3466
        2.0         2917
        1.0         1395
        4.0         1143
        5.0           33
        6.0            3
dtype: int64
NBF Data        
 CBF_01  P_NUMFLU
2       0.0         656
        3.0         503
        2.0         422
        1.0         277
        4.0         137
        5.0           2
dtype: int64
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
0    3466    0
1    1395    1395
2    2917    5834
3    4334    13002
4    1143    4572
5    33    165
6    3    18
BF Data        
 CBF_01  P_NUMFLU
1       3.0         4334
        0.0         3466
        2.0         2917
        1.0         1395
        4.0         1143
        5.0           33
        6.0            3
dtype: int64
NBF Data        
 CBF_01  P_NUMFLU
2       0.0         656
        3.0         503
        2.0         422
        1.0         277
        4.0         137
        5.0           2
dtype: int64
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    columns                 = ['SEX', 'HAD_CPOX', 'P_NUMVRC']
    vaccine_data_df         = pd.read_csv("NISPUF17.csv", usecols=columns)
    df_len                  = len(vaccine_data_df)

    # get entries with 'P_NUMVRC' = 1, we only need vaccinated individuals
    vaccinated_df     = vaccine_data_df[vaccine_data_df['P_NUMVRC'] > 0]
    
    gender_df         = vaccinated_df[vaccinated_df['SEX'] <= 2]
    had_pox_df        = vaccinated_df[vaccinated_df['HAD_CPOX'] <= 2]
    gender_had_pox_df = had_pox_df[had_pox_df['SEX'] <= 2] 

    # Series containing 
    vaccinated_data = vaccine_data_df.value_counts()
    gender_data = gender_df.value_counts()
    had_pox_data = had_pox_df.value_counts()

    boy_vaccinated_without_cpox  = 0     # 1, 1, > 0
    boy_vaccinated_with_cpox     = 0     # 1, 2, > 0

    girl_vaccinated_without_cpox = 0     # 2, 1, > 0
    girl_vaccinated_with_cpox    = 0     # 2, 2, > 0

    for gender in range(1, 3) :
        for had_pox in range(1, 3) :
            for n_vaccines in range(1, len(vaccinated_data)+1) :
                try :
                    had_pox_data[had_pox, gender, n_vaccines]
                except KeyError :
                    continue

                if gender == 1 :
                    if had_pox == 1 :
                        boy_vaccinated_without_cpox += had_pox_data[had_pox, gender, n_vaccines]
                    else :
                        boy_vaccinated_with_cpox += had_pox_data[had_pox, gender, n_vaccines]
                else:
                    if had_pox == 1 :
                        girl_vaccinated_without_cpox += had_pox_data[had_pox, gender, n_vaccines]
                    else :
                        girl_vaccinated_with_cpox += had_pox_data[had_pox, gender, n_vaccines]

    print ("boy_vaccinated_with_cpox", boy_vaccinated_with_cpox)
    print ("boy_vaccinated_without_cpox", boy_vaccinated_without_cpox)
    print ("girl_vaccinated_without_cpox", girl_vaccinated_without_cpox)
    print ("girl_vaccinated_with_cpox", girl_vaccinated_with_cpox)
    if boy_vaccinated_with_cpox == 0 :
        raise NotImplementedError()
    if girl_vaccinated_with_cpox == 0 :
        raise NotImplementedError()

    return {"male": boy_vaccinated_without_cpox/boy_vaccinated_with_cpox, "female": girl_vaccinated_without_cpox/girl_vaccinated_with_cpox}

    raise NotImplementedError()
```

```python
print(chickenpox_by_sex())
```

Output:
```
boy_vaccinated_with_cpox 7028
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
#                   "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
#    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    columns                 = ['HAD_CPOX', 'P_NUMVRC']
    vaccine_data_df         = pd.read_csv("NISPUF17.csv", usecols=columns)
    df_len                  = len(vaccine_data_df)
    
    vaccine_data_df.dropna(axis=0, inplace=True)

    clean_df=pd.DataFrame({"had_chickenpox_column":vaccine_data_df['HAD_CPOX'].to_numpy(),
                   "num_chickenpox_vaccine_column":vaccine_data_df['P_NUMVRC'].to_numpy()})


#    corr = 0.;
    corr, pval=stats.pearsonr(clean_df["had_chickenpox_column"],clean_df["num_chickenpox_vaccine_column"])
    print(corr, ' ', pval)

    return corr

    raise NotImplementedError()
```

```python
corr_chickenpox()
# 0.0034753171515293467   0.6669755758983259
```

Output:
```
0.003475317151529321   0.6669755758989325
0.003475317151529321
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
0.003475317151529321   0.6669755758989325
```

## FINAL SUBMISSION 100%

```python
import pandas as pd
import numpy as np

def proportion_of_education():

    # Maternal education 
    # 1: <12 years, 
    # 2:12 years, 
    # 3: >12 years, not a college graduate
    # 4: College graduate
    columns = ['EDUC1']
    education_df = pd.read_csv("assets/NISPUF17.csv", usecols=columns)
    df_len = len(education_df)
    data = education_df['EDUC1'].value_counts()
    
  #  print({"less than high school": data.get(1.0)/df_len, "high school": data.get(2.0)/df_len, "more than high school but not college": data.get(3.0)/df_len, "college": data.get(4.0)/df_len})
    return {"less than high school": data.get(1.0)/df_len, "high school": data.get(2.0)/df_len, "more than high school but not college": data.get(3.0)/df_len, "college": data.get(4.0)/df_len}

    raise NotImplementedError()
```

```python
def average_influenza_doses():
    # Was child ever fed breast milk?
    #    Yes
    #    No
    #    Don’t know
    #    Missing
    
    columns               = ['CBF_01', 'P_NUMFLU']
    bf_and_vaccine_df     = pd.read_csv("assets/NISPUF17.csv", usecols=columns)
    df_len                = len(bf_and_vaccine_df)

    # remove NA data from breast_and_vaccine_df
    bf_and_vaccine_df.dropna(axis=0, inplace=True)

    # get entries with CBF_01 == 1 AND CBF_01 == 2
    breast_fed     = bf_and_vaccine_df[bf_and_vaccine_df['CBF_01'] == 1]
    not_breast_fed = bf_and_vaccine_df[bf_and_vaccine_df['CBF_01'] == 2]

    # get counts of breast fed and 
    bf_data  = breast_fed['P_NUMFLU'].value_counts()
    nbf_data = not_breast_fed['P_NUMFLU'].value_counts()

    n_bf_injected          = 0
    n_bf_total_injections  = 0
    for i in range(len(bf_data)) :
        n_bf_injected         = n_bf_injected + bf_data[i]
        n_bf_total_injections = n_bf_total_injections + (i * bf_data[i])

    n_nbf_injected         = 0
    n_nbf_total_injections = 0
    for i in range(len(nbf_data)) :
        n_nbf_injected         = n_nbf_injected + nbf_data[i]
        n_nbf_total_injections = n_nbf_total_injections + (i * nbf_data[i])

    return (n_bf_total_injections/n_bf_injected, n_nbf_total_injections/n_nbf_injected)
    raise NotImplementedError()
```

```python
import pandas as pd
import numpy as np

def chickenpox_by_sex():

    columns                 = ['SEX', 'HAD_CPOX', 'P_NUMVRC']
    vaccine_data_df         = pd.read_csv("assets/NISPUF17.csv", usecols=columns)
    df_len                  = len(vaccine_data_df)

    # get entries with 'P_NUMVRC' = 1, we only need vaccinated individuals
    clean_vaccine_data_df = vaccine_data_df.dropna(axis=0, inplace=False)

    # filter vaccinated individuals only
    vaccinated_df                       = clean_vaccine_data_df[clean_vaccine_data_df['P_NUMVRC'] > 0]

    # filter vaccinated individuals by gender
    vaccinated_male_df                  = vaccinated_df[vaccinated_df['SEX'] == 1]
    vaccinated_female_df                = vaccinated_df[vaccinated_df['SEX'] == 2]

    # filter vaccinated individuals by gender and if they have had CPOX
    n_vaccinated_male_cpox_yes         = len(vaccinated_male_df[vaccinated_male_df['HAD_CPOX'] == 1])
    n_vaccinated_male_cpox_no          = len(vaccinated_male_df[vaccinated_male_df['HAD_CPOX'] == 2])
    n_vaccinated_female_cpox_yes       = len(vaccinated_female_df[vaccinated_female_df['HAD_CPOX'] == 1])
    n_vaccinated_female_cpox_no        = len(vaccinated_female_df[vaccinated_female_df['HAD_CPOX'] == 2])

    male_ratio                         = n_vaccinated_male_cpox_yes/n_vaccinated_male_cpox_no
    female_ratio                       = n_vaccinated_female_cpox_yes/n_vaccinated_female_cpox_no    

    return {"male": male_ratio, "female": female_ratio}

    raise NotImplementedError()

```

```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    columns                 = ['HAD_CPOX', 'P_NUMVRC']
    vaccine_data_df         = pd.read_csv("assets/NISPUF17.csv", usecols=columns)
    df_len                  = len(vaccine_data_df)
    
    clean_vaccine_data_df   = vaccine_data_df.dropna(axis=0, inplace=False)
    clean_vaccine_data_df   = clean_vaccine_data_df[clean_vaccine_data_df.HAD_CPOX < 3]
    
    cpox_column    = clean_vaccine_data_df['HAD_CPOX'].to_numpy()
    vaccine_column = clean_vaccine_data_df['P_NUMVRC'].to_numpy()
    
    
    clean_df=pd.DataFrame({"had_chickenpox_column":cpox_column,
                   "num_chickenpox_vaccine_column":vaccine_column})

    corr, pval=stats.pearsonr(clean_df["had_chickenpox_column"],clean_df["num_chickenpox_vaccine_column"])

    return corr

    raise NotImplementedError()   
```