# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import sys
#sys.executable
!pip install --user openpyxl
```

Output:
```
Requirement already satisfied: openpyxl in /home/kevin/.local/lib/python3.8/site-packages (3.0.9)
Requirement already satisfied: et-xmlfile in /home/kevin/.local/lib/python3.8/site-packages (from openpyxl) (1.1.0)
```

```python
import pandas as pd
import numpy as np
import re
import openpyxl
import unicodedata

pd.set_option('display.max_rows', 500)

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
#import warnings
#warnings.filterwarnings('ignore')
```

```python

```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python

```

## def getEnergyIndicators(excelName)
filter the directories as required 
  1) exclude the footer and header information from the datafile. 
  2) The first two columns are unneccessary, so you should get rid of them, 
  3) change the column labels so that the columns are:
        ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]
  4) Convert Energy Supply to gigajoules (Note: there are 1,000,000 
        gigajoules in a petajoule). 
  5) For all countries which have missing data (e.g. data with '...') make
     sure this is reflected as np.NaN values.
  6) Rename the following list of countries (for use in later questions):
         'Republic of Korea': 'South Korea',
         'United States of America': 'United States',
         'United Kingdom of Great Britain and Northern Ireland': 'United Kingdom',
         'China, Hong Kong Special Administrative Region': 'Hong Kong'
  7) There are also several countries with numbers and/or parenthesis in 
     their name. Be sure to remove these, 
          e.g. 'Bolivia (Plurinational State of)' should be 'Bolivia'. 
               'Switzerland17' should be 'Switzerland'.


```python

```

```python
def getEnergyIndicators() :

    # Requirements 1, 2, 5
    energyIndicators   = pd.read_excel(
        io             = 'assets/Energy Indicators.xls',
        sheet_name     = 'Energy',
        na_values      = '...',
        usecols        = 'C:F',
        skiprows       = range(17),
        nrows          = 247-20
    )
    # Requirements 3, 5
    energyIndicators.columns    = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']

    # Requirements 7
    energyIndicators['Country'] = energyIndicators['Country'].str.replace('\d+', '',case=True, regex=True)
    energyIndicators['Country'] = energyIndicators['Country'].str.replace(r' \([^)]*\)', '', regex=True)
    energyIndicators['Country'] = energyIndicators['Country'].where(energyIndicators['Country'].notna(), np.NaN)

    # Set index to the country column
    energyIndicators.set_index('Country', inplace=True)

    # Requirements 6
    energyIndicators.rename(index={
        'Republic of Korea': 'South Korea', 
        'United States of America': 'United States',
        'United Kingdom of Great Britain and Northern Ireland': 'United Kingdom',
        'China, Hong Kong Special Administrative Region': 'Hong Kong'
    }, inplace=True)

    # Requirements 4
    energyIndicators['Energy Supply'] = energyIndicators['Energy Supply'].apply(lambda x: x * 1e6)

    return energyIndicators

    raise NotImplementedError()

```

```python
getEnergyIndicators()
```

Output:
```
Energy Supply  \
Country                                                     
Afghanistan                                  3.210000e+08   
Albania                                      1.020000e+08   
Algeria                                      1.959000e+09   
American Samoa                                        NaN   
Andorra                                      9.000000e+06   
Angola                                       6.420000e+08   
Anguilla                                     2.000000e+06   
Antigua and Barbuda                          8.000000e+06   
Argentina                                    3.378000e+09   
Armenia                                      1.430000e+08   
Aruba                                        1.200000e+07   
Australia                                    5.386000e+09   
Austria                                      1.391000e+09   
Azerbaijan                                   5.670000e+08   
Bahamas                                      4.500000e+07   
Bahrain
```

## def getGDP(csvName) 
filter the directories as required 
1. Make sure to skip the header, 
2. rename the following list of countries:
        "Korea, Rep.": "South Korea", 
        "Iran, Islamic Rep.": "Iran",
        "Hong Kong SAR, China": "Hong Kong"
3. Use only the last 10 years (2006-2015) of GDP data

```python
def getGDP() :
    # Read 'Country Name', 2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', 2015'
    
    colList = ['Country Name', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']
    gdp = pd.read_csv(
        'assets/world_bank.csv', 
        header = 4,
        usecols=colList
    )
    gdp.columns = ['Country', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']

    gdp_obj = gdp.select_dtypes(['object'])
    gdp[gdp_obj.columns] = gdp_obj.apply(lambda x: x.str.strip())

    gdp.set_index('Country', inplace=True)

    gdp.rename(index={
        'Korea, Rep.': 'South Korea', 
        'Iran, Islamic Rep.': 'Iran',
        'Hong Kong SAR, China': 'Hong Kong'
    }, inplace=True)

    return gdp

    raise NotImplementedError()

```

```python
print(getGDP())
```

Output:
```
2006  \
Country                                                            
Aruba                                                        NaN   
Andorra                                             4.018196e+09   
Afghanistan                                         1.030523e+10   
Angola                                              5.581103e+10   
Albania                                             9.771760e+09   
Arab World                                          1.763824e+12   
United Arab Emirates                                2.789489e+11   
Argentina                                           3.788140e+11   
Armenia                                             8.679746e+09   
American Samoa                                               NaN   
Antigua and Barbuda                                 1.268714e+09   
Australia                                           1.021939e+12   
Austria                                             3.782096e+11   
Azerbaijan                              
```

## Finally, load the Sciamgo Journal and Country Rank data for Energy Engineering and 
 Power Technology from the file assets/scimagojr-3.xlsx, 
 which ranks countries based on their journal contributions in the aforementioned area. 
 Call this DataFrame ScimEn.

 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index'

```python
def getScienceRank() :
    scienceRank = pd.read_excel(
        'assets/scimagojr-3.xlsx',
        sheet_name='Sheet1',
        usecols='A:H'
    )
    scienceRank_obj = scienceRank.select_dtypes(['object'])
    scienceRank[scienceRank_obj.columns] = scienceRank_obj.apply(lambda x: x.str.strip())
    scienceRank['Country'].apply(lambda val: unicodedata.normalize('NFD', val).encode('ascii', 'ignore').decode())

    scienceRank.columns = ['Rank', 'Country', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index']
    scienceRank.set_index('Country', inplace=True)

    return scienceRank

    raise NotImplementedError()

```

```python
print (getScienceRank())
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                                         
China                              1     127050             126767     597237   
United States                      2      96661              94747     792274   
Japan                              3      30504              30287     223024   
United Kingdom                     4      20944              20357     206091   
Russian Federation                 5      18534              18301      34266   
Canada                             6      17899              17620     215003   
Germany                            7      17027              16831     140566   
India                              8      15005              14841     128763   
France                             9      13153              12973     130632   
South Korea                       10      11983              11923     114675   
Italy                             11      10964             
```

## Final output  

      ['Rank', 
      'Documents', 
      'Citable documents', 
      'Citations', 
      'Self-citations', 
      'Citations per document', 
      'H index', 
      'Energy Supply', 
      'Energy Supply per Capita', 
      '% Renewable', 
      '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'].
Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15).

The index of this DataFrame should be the name of the country, and the columns should be \['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per Capita', '\% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'\].

```python
def mergedCompleteDataframe():
    Energy           = getEnergyIndicators()
    GDP              = getGDP()
    ScimEn           = getScienceRank()

    # join dataframe
    mergedCompleteDF = pd.merge(pd.merge(ScimEn,Energy,on='Country', how = 'inner'),GDP,on='Country', how = 'inner')    
#    mergedCompleteDF.set_index('Country', inplace=True)    

#    mergedDF = pd.merge(ScimEn,Energy,on='Country')
#    print("MERGED \n", mergedDF)
#    mergedCompleteDF = pd.merge(mergedDF,GDP,on='Country')
#    print("MERGED COMPLETE \n", mergedCompleteDF)

    return mergedCompleteDF

    raise NotImplementedError()

```

```python
mergedDF = mergedCompleteDataframe()

print("merged length = ", len(mergedDF))
print(mergedDF.index.name)
print(mergedDF.info())
print(mergedDF)
```

Output:
```
merged length =  162
Country
<class 'pandas.core.frame.DataFrame'>
Index: 162 entries, China to Mauritania
Data columns (total 20 columns):
 #   Column                    Non-Null Count  Dtype  
---  ------                    --------------  -----  
 0   Rank                      162 non-null    int64  
 1   Documents                 162 non-null    int64  
 2   Citable documents         162 non-null    int64  
 3   Citations                 162 non-null    int64  
 4   Self-citations            162 non-null    int64  
 5   Citations per document    162 non-null    float64
 6   H index                   162 non-null    int64  
 7   Energy Supply             161 non-null    float64
 8   Energy Supply per Capita  161 non-null    float64
 9   % Renewable               162 non-null    float64
 10  2006                      153 non-null    float64
 11  2007                      153 non-null    float64
 12  2008                      153 non-null    float64
 13  2009                      153 
```

      \['Rank', 
      'Documents', 
      'Citable documents', 
      'Citations', 
      'Self-citations', 
      'Citations per document', 
      'H index', 
      'Energy Supply', 
      'Energy Supply per Capita', 
      '\% Renewable', 
      '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'\].

```python
print(type(loadAnalysisFrame()))
print(loadAnalysisFrame())


```

```python
def loadAnalysisFrame():
    analysisFrame    = mergedCompleteDataframe()

    return analysisFrame.head(15)

    raise NotImplementedError()

```

```python
def answer_one():
    df  = loadAnalysisFrame()

    return df

    raise NotImplementedError()
```

```python
print(answer_one())
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Iran                  13       8896               8819
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

```python

```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    
    # wrong approach do an outer merge followed by an inner merge and take the diff, just read the damn instructions   
    #    return len(mergedCompleteDataframe()) - len(loadAnalysisFrame())
    Energy           = getEnergyIndicators()
    GDP              = getGDP()
    ScimEn           = getScienceRank()

    innerJoin = mergedCompleteDataframe()
    outerJoin = pd.merge(pd.merge(ScimEn,Energy,on='Country', how = 'outer'),GDP,on='Country', how = 'outer')

    print(len(innerJoin))
    print(len(outerJoin))
    print(len(outerJoin) - len(innerJoin))

    return len(outerJoin) - len(innerJoin)

    raise NotImplementedError()
```

```python
print(type(answer_two()))
print(answer_two())
```

Output:
```
162
318
156
<class 'int'>
162
318
156
156
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():

    avgGDP    = loadAnalysisFrame()
    colList   = ['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']
    avgGDP    = avgGDP[colList]

    # compute average GDP for the items in the analysisFrame on each row from 2006-2015
    avgGDP['avgGDP'] = avgGDP.mean(axis=1, skipna=True, numeric_only=True)

    # sort avgGDP highest to lowest
    avgGDP.sort_values(by=['avgGDP'], axis=0, ascending=False, inplace=True, kind='quicksort', na_position='last')

    return avgGDP['avgGDP'].head(15)

    raise NotImplementedError()

```

```python
ser = answer_three()
print(len(ser))
print(type(ser))
print(ser)
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():

    gdp       = loadAnalysisFrame()
    colList   = ['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']
    gdp       = gdp[colList]
    
    name      = gdp.index[5]
    print(name)

    print(gdp['2006'][name])
    print(gdp['2015'][name])
    print(gdp['2015'][name] - gdp['2006'][name])

    return gdp['2015'][name] - gdp['2006'][name]

    raise NotImplementedError()
```

```python
span = answer_four()
print(span)
print(type(span))
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    energySupplyPerHead       = loadAnalysisFrame()
    
    return energySupplyPerHead['Energy Supply per Capita'].mean(axis=0, skipna=True)
    raise NotImplementedError()
```

```python
avgEnergySupply = answer_five()
print(type(avgEnergySupply))
print(avgEnergySupply)

```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():

    maxRenewableDF   = loadAnalysisFrame()
    renewable        = maxRenewableDF['% Renewable']
    maxRenewable     = renewable.max()
    country          = renewable.idxmax()

    return (country, maxRenewable)

    raise NotImplementedError()
```

```python
print(answer_six())
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():

    citationDF      = loadAnalysisFrame()
    citationDF['ratio'] = citationDF['Self-citations']/citationDF['Citations']

    maxRatio         = citationDF['ratio'].max()
    country          = citationDF['ratio'].idxmax()

    return (country, maxRatio)

    raise NotImplementedError()
```

```python
print(answer_seven())
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

```python
def populationEstimate():
    populationDF        = loadAnalysisFrame()
    populationDF['popEst'] = populationDF['Energy Supply']/populationDF['Energy Supply per Capita']
    populationDF.sort_values(by=['popEst'], axis=0, ascending=False, inplace=True, kind='quicksort', na_position='last')

    return populationDF
```

```python
populationEstimate()['popEst']
```

Output:
```
Country
China                 1.367645e+09
India                 1.276731e+09
United States         3.176154e+08
Brazil                2.059153e+08
Russian Federation    1.435000e+08
Japan                 1.274094e+08
Germany               8.036970e+07
Iran                  7.707563e+07
United Kingdom        6.387097e+07
France                6.383735e+07
Italy                 5.990826e+07
South Korea           4.980543e+07
Spain                 4.644340e+07
Canada                3.523986e+07
Australia             2.331602e+07
Name: popEst, dtype: float64
```

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    populationDF = populationEstimate()
    return populationDF.index[6]

    raise NotImplementedError()
```

```python
print(answer_eight())
```

Output:
```
Germany
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():

    populationDF = populationEstimate()
    dataDF       = loadAnalysisFrame()
    dataDF['Citable Documents per Person']  = dataDF['Citable documents']/populationDF['popEst']

    correlation = dataDF.filter(['Energy Supply per Capita', 'Citable Documents per Person'], axis=1).corr()
    
    return correlation['Citable Documents per Person']['Energy Supply per Capita']
    raise NotImplementedError()
```

```python
print(answer_nine())
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    renewableDF                 = loadAnalysisFrame()
    renewableDF['Renewable']    = renewableDF['% Renewable']
    renewableMedian             = renewableDF['% Renewable'].median(axis=0, skipna=True)

    renewableDF.loc[(renewableDF['% Renewable'] >= renewableMedian),'Renewable'] = 1
    renewableDF.loc[(renewableDF['% Renewable'] < renewableMedian),'Renewable'] = 0
    renewableDF.sort_values(by=['Rank'], axis=0, ascending=True, inplace=True, kind='quicksort', na_position='last')
    renewableDF['Renewable'] = renewableDF['Renewable'].astype(int)

    HighRenew = renewableDF['Renewable']
#    print(HighRenew.index)

    return HighRenew

    raise NotImplementedError()
```

```python
highRenew = answer_ten()
print(len(highRenew))
print(highRenew)

```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

```python

#     sample code for iteratively accessing continent/country data
#     for continent in range(len(ContinentData)):
#         print('continent = ', ContinentData.iloc[continent], ' type = ', type(ContinentData.iloc[continent][0]))
#         print('country list = ', ContinentData.iloc[continent][0])
#         for country in ContinentData.iloc[continent]['ContinentCountries']:
#             print('country = ', country, ' type = ', type(country))
# continent =  ContinentCountries    [China, Japan, India, South Korea, Iran]
# Name: Asia, dtype: object  type =  <class 'list'>
# country list =  ['China', 'Japan', 'India', 'South Korea', 'Iran']
# country =  China  type =  <class 'str'>
# country =  Japan  type =  <class 'str'>
# country =  India  type =  <class 'str'>
# country =  South Korea  type =  <class 'str'>
# country =  Iran  type =  <class 'str'>
# continent =  ContinentCountries    [United States, Canada]
# Name: North America, dtype: object  type =  <class 'list'>
# country list =  ['United States', 'Canada']
# country =  United States  type =  <class 'str'>
# country =  Canada  type =  <class 'str'>
# continent =  ContinentCountries    [United Kingdom, Russian Federation, Germany, ...
# Name: Europe, dtype: object  type =  <class 'list'>
# country list =  ['United Kingdom', 'Russian Federation', 'Germany', 'France', 'Italy', 'Spain']
# country =  United Kingdom  type =  <class 'str'>
# country =  Russian Federation  type =  <class 'str'>
# country =  Germany  type =  <class 'str'>
# country =  France  type =  <class 'str'>
# country =  Italy  type =  <class 'str'>
# country =  Spain  type =  <class 'str'>
# continent =  ContinentCountries    [Australia]
# Name: Australia, dtype: object  type =  <class 'list'>
# country list =  ['Australia']
# country =  Australia  type =  <class 'str'>
# continent =  ContinentCountries    [Brazil]
# Name: South America, dtype: object  type =  <class 'list'>
# country list =  ['Brazil']
# country =  Brazil  type =  <class 'str'>


    # create a list of continent country lists
#     continentCountryList = []
#     for continent in range(len(ContinentData)):
#         continentCountryList.append(ContinentData.iloc[continent][0])

    
    # load the population data for the required countries
#     populationData = populationEstimate()['popEst']
#    print(populationData)
#     continentCountryPop = populationData[continentCountryList[0]]
#    print('cont country pop  = ', continentCountryPop, ' sum = ',continentCountryPop.sum())

#     continentList = list(ContinentData.index.values)

#     measureListDictionary = {measure[0] : [], measure[1] : [], measure[2] : [], measure[3] : []}
    
#     for continent in range(len(continentList)):
#         continentCountryPop = populationData[continentCountryList[continent]]
#        print('cont country pop  = ', continentCountryPop, ' sum = ',continentCountryPop.sum())
#        print('name = ', name, ' datalist = ',dataList)
#        print(measureListDictionary['mean'])
#         measureListDictionary[measure[0]].append(len(continentCountryPop))
#         measureListDictionary[measure[1]].append(continentCountryPop.sum())
#         measureListDictionary[measure[2]].append(continentCountryPop.mean())
#         measureListDictionary[measure[3]].append(continentCountryPop.std())

#     for idx in measure:
#         dataFrame = pd.DataFrame(measureListDictionary[idx], index=continentList)
##        print('dataFrame = ', dataFrame)
#         dataFrame.index.name='Continent'
#         ContinentData = pd.merge(ContinentData, dataFrame, on='Continent')
##    print('ContinentData = ', ContinentData)

#     ContinentData.drop(['ContinentCountries'], axis=1, inplace=True)
#     ContinentData.columns = ['size', 'sum', 'mean', 'std']
#     ContinentData['mean'] = ContinentData['mean'].astype(float)
    #concatenate all series

#    print(ContinentData, "\n", type(ContinentData))

```

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
# Target: a dataframe containing a list of continents and an index
#         a list of countries on the continent
#         a list of country populations on the continent
#         a list of the number of countries/continent
#         a list of sum, mean, and std for each continent
#     
# Create a ContinentData dataframe consisting of the following lists
#     1) List of continents
#     2) A list of lists of countries
#     3) A list of lists of the population in each country
def answer_eleven():
    ContinentDict  = {
        'China':'Asia', 
        'United States':'North America', 
        'Japan':'Asia', 
        'United Kingdom':'Europe', 
        'Russian Federation':'Europe', 
        'Canada':'North America', 
        'Germany':'Europe', 
        'India':'Asia',
        'France':'Europe', 
        'South Korea':'Asia', 
        'Italy':'Europe', 
        'Spain':'Europe', 
        'Iran':'Asia',
        'Australia':'Australia', 
        'Brazil':'South America'
    }
    

    # create ContinentData from ContinentDict and label index and column
    ContinentCountries = pd.Series(dtype='str')
    for i, v in ContinentDict.items():
        ContinentCountries[v] = [i] if v not in ContinentCountries.keys() else ContinentCountries[v] + [i]

    ContinentData = pd.DataFrame(ContinentCountries)
    ContinentData.reset_index(inplace=True)
    ContinentData.columns =['Continent','ContinentCountries']

    population = populationEstimate()['popEst']

    continentPop  = []
    continentSize = []
    continentSum  = []
    continentMean = []
    continentStd  = []
    for continent in ContinentData['ContinentCountries']:
        pop = []
        for country in continent:
            pop.append(population[country].astype(int))
        continentPop.append(pop)
        continentSize.append(len(pop))
        continentSum.append(sum(pop))
        continentMean.append(np.mean(pop))
        continentStd.append(np.std(pop))

    ContinentData['Population'] = continentPop    
    ContinentData['size'] = continentSize    
    ContinentData['sum'] = continentSum    
    ContinentData['mean'] = continentSum    
    ContinentData['std'] = continentSum    

    ContinentData.index = ContinentData['Continent']

    ContinentData.drop(['Continent'], axis=1, inplace=True)
    ContinentData.drop(['ContinentCountries'], axis=1, inplace=True)
    ContinentData.drop(['Population'], axis=1, inplace=True)

    return ContinentData
    raise NotImplementedError()
    

```

```python
continentData = answer_eleven()
print(continentData)
print(type(continentData))
```

Output:
```
size         sum        mean         std
Continent                                              
Asia              5  2898666384  2898666384  2898666384
North America     2   352855248   352855248   352855248
Europe            6   457929664   457929664   457929664
Australia         1    23316017    23316017    23316017
South America     1   205915254   205915254   205915254
<class 'pandas.core.frame.DataFrame'>
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```